# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).<br>

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays

<br>
As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). Dictionaries and arrays are both mutable.
We'll explain this more below!

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Output:
```
("penguins", "cats", "sugargliders")
```

We can index into this tuple,

```julia
myfavoriteanimals[0]
```

Output:
```
@ Base ./namedtuple.jl:167
 [2] top-level scope
   @ In[12]:1
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[11]:1
```

## Now in 1.0: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)
```

A good example is a contacts list, where we associate names with phone numbers.

```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

In this example, each name and number is a "key" and "value" pair. We can grab Jenny's number (a value) using the associated key

```julia
myphonebook["Jenny"]
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Kramer"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Let's check what our phonebook looks like now...

```julia
myphonebook
```

We can delete Kramer from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Kramer")
```

Output:
```
"555-FILK"
```

`function!` means the function will change input argument.

`function` means the function will return something without changing input arguments.

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
myphonebook[1]
```

Output:
```
@ Base ./dict.jl:498
 [2] top-level scope
   @ In[9]:1
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

## Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```


For example, we might create an array to keep track of my friends

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]
```

The `1` in `Array{String,1}` means this is a one dimensional vector.  An `Array{String,2}` would be a 2d matrix, etc.  The `String` is the type of each element.

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third friend listed in `myfriends`, we write

```julia
myfriends[3]
```

We can use indexing to edit an existing element of an array

```julia
myfriends[3] = "Baby Bop"
```

Yes, Julia is 1-based indexing, not 0-based like Python.  Wars are fought over lesser issues. I have a friend with the wisdom of Solomon who proposes settling this once and for all with ½ 😃

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

and then remove it

```julia
pop!(fibonacci)
```

```julia
fibonacci
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Below are examples of 2D and 3D arrays populated with random values.

```julia
rand(4, 3)
```

```julia
rand(4, 3, 2)
```

Be careful when you want to copy arrays!

```julia
fibonacci
```

```julia
somenumbers = fibonacci
```

```julia
somenumbers[1] = 404
```

```julia
fibonacci
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

```julia
somemorenumbers = copy(fibonacci)
```

```julia
somemorenumbers[1] = 404
```

```julia
fibonacci
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

### Exercises

#### 3.1 
Create an array, `a_ray`, with the following code:

```julia
a_ray = [1, 2, 3]

```

Add the number `4` to the end of this array and then remove it.

```julia
a_ray = [1, 2, 3];
push!(a_ray, 4)
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
myphonebook["Emergency"] = 911  # dict{A=>B} A and B types are defined.
```

Output:
```
@ Base ./dict.jl:376
 [2] top-level scope
   @ In[20]:1
```

```julia
myphonebook["Emergency"] = string(911);
myphonebook["Emergency"]
```

Output:
```
"911"
```

#### 3.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")

```

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368") # but can be defined as "ANY"
```

Output:
```
Dict{String, Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia
flexible_phonebook["Emergency"] = 911
```

Output:
```
911
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values? (hint: try using [Julia's documentation for dictionaries](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

```julia
?dict
```

Output:
```
search: [0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m I[0m[1md[22mD[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m WeakKey[0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m Abstract[0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m re[0m[1md[22m[0m[1mi[22mre[0m[1mc[22m[0m[1mt[22m_stdio re[0m[1md[22m[0m[1mi[22mre[0m[1mc[22m[0m[1mt[22m_stdin

Couldn't find [36mdict[39m
Perhaps you meant Dict, IdDict, diff, div, edit, Cint, Dims, digits or fdio
  No documentation found.

  Binding [36mdict[39m does not exist.
```

```julia
?Dict # ModuleName, TypeName, functionanme, variablename, macroname, verylong_functionname
```

Output:
```
search: [0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m I[0m[1md[22mD[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m WeakKey[0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m Abstract[0m[1mD[22m[0m[1mi[22m[0m[1mc[22m[0m[1mt[22m re[0m[1md[22m[0m[1mi[22mre[0m[1mc[22m[0m[1mt[22m_stdio re[0m[1md[22m[0m[1mi[22mre[0m[1mc[22m[0m[1mt[22m_stdin

[36m  Dict([itr])[39m

  [36mDict{K,V}()[39m constructs a hash table with keys of type [36mK[39m and values of type
  [36mV[39m. Keys are compared with [36misequal[39m and hashed with [36mhash[39m.

  Given a single iterable argument, constructs a [36mDict[39m whose key-value pairs
  are taken from 2-tuples [36m(key,value)[39m generated by the argument.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡[22m

[36m  julia> Dict([("A", 1), ("B", 2)])[39m
[36m  Dict{String, Int64} with 2 entries:[39m
[36m    "B" => 2[39m
[36m    "A" => 1[39m

  Alternatively, a sequence of pair arguments may be pass
```

Please click on `Validate` button at the top, upon completion of the exercise